# Problemas

---
Crea una función que sea capaz de calcular la matriz de rotación en un plano con cualquier angulo que se le de.

<!-- grade_id: cell-cbe6b534af68b731, points: 0, grade: false, solution: true -->
```python
def rotacion(θ):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3767258f4db444e1, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import matrix
assert_allclose(rotacion(0), matrix([[1, 0], [0, 1]]), rtol=1e-5, atol=1e-5)
```

---
Calcule la matriz de rotación en un plano, con un angulo de 30º, 50º y 100º, y guardalos en las variables ```A1```, ```A2```, ```A3```.

<!-- grade_id: cell-5856c502c0bdac7c, points: 0, grade: false, solution: true -->
```python
from numpy import pi
τ = 2*pi
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c84ab72c9098ebac, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
```

---
Aplique las matrices de rotación de 30º, 100º y 50º al vector $\begin{pmatrix} 2 \\ 2\end{pmatrix}$ y guarde el resultado en la variable ```res1```.

<!-- grade_id: cell-44ddf2f4d6904fa6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-54cd99b9320ea064, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
```

---
Aplique las matrices de rotación de 50º, 100º y 30º al vector $\begin{pmatrix} 2 \\ 2\end{pmatrix}$ y guarde el resultado en la variable ```res2```.

<!-- grade_id: cell-1f99db1538761d2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7423316a8b222d0c, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
```