# Problemas

---
Defina una función que tome como argumentos los parametros de Denavit - Hartenberg, y cree una matriz de transformación homogénea.

<!-- grade_id: cell-d796d71f88ca3f1a, points: 0, grade: false, solution: true -->
```python
def DH_simbolico(a, d, α, θ):
    from sympy import Matrix, sin, cos
    ### BEGIN SOLUTION
    A = Matrix([[cos(θ), -sin(θ)*cos(α), sin(θ)*sin(α), a*cos(θ)],
                 [sin(θ), cos(θ)*cos(α), -cos(θ)*sin(α), a*sin(θ)],
                 [0, sin(α), cos(α), d],
                 [0, 0, 0, 1]])
    return A
    ### END SOLUTION
```

<!-- grade_id: cell-2e02d67eabed1e86, points: 2, grade: true, solution: false -->
```python
from sympy import Matrix, sin, cos, pi
from nose.tools import assert_equal
assert_equal(DH_simbolico(0,0,0,pi/2), Matrix([[0,-1,0,0],[1,0,0,0], [0,0,1,0],[0,0,0,1]]))
assert_equal(DH_simbolico(0,0,pi/2,0), Matrix([[1,0,0,0],[0,0,-1,0], [0,1,0,0],[0,0,0,1]]))
assert_equal(DH_simbolico(0,1,0,0), Matrix([[1,0,0,0],[0,1,0,0], [0,0,1,1],[0,0,0,1]]))
assert_equal(DH_simbolico(1,0,0,0), Matrix([[1,0,0,1],[0,1,0,0], [0,0,1,0],[0,0,0,1]]))
```

---
Cree una función que tome como argumentos los parametros de los grados de libertad de un manipulador tipo PUMA y devuelva las matrices de transformación homogénea asociadas a cada eslabon.

<!-- grade_id: cell-8759bf18b64c88c0, points: 0, grade: false, solution: true -->
```python
def cinematica_PUMA(q1, q2, q3):
    from sympy import pi, var
    var("l1:4")
    ### BEGIN SOLUTION
    A1 = DH_simbolico(0,l1,pi/2, q1)
    A2 = DH_simbolico(l2, 0, 0, q2)
    A3 = DH_simbolico(l3, 0, 0, q3)
    ### END SOLUTION
    return A1, A2, A3
```

<!-- grade_id: cell-5bdcfe97ca2cef34, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
from sympy import pi, var, Matrix
var("l1:4")
A1, A2, A3 = cinematica_PUMA(0, 0, 0)
assert_equal(A1*A2*A3, Matrix([[1,0,0,l2+l3], [0,0,-1,0], [0,1,0,l1], [0,0,0,1]]))
A1, A2, A3 = cinematica_PUMA(pi/2, 0, 0)
assert_equal(A1*A2*A3, Matrix([[0,0,1,0], [1,0,0,l2+l3], [0,1,0,l1], [0,0,0,1]]))
A1, A2, A3 = cinematica_PUMA(0, pi/2, 0)
assert_equal(A1*A2*A3, Matrix([[0,-1,0,0], [0,0,-1,0], [1,0,0,l1+l2+l3], [0,0,0,1]]))
A1, A2, A3 = cinematica_PUMA(0, 0, pi/2)
assert_equal(A1*A2*A3, Matrix([[0,-1,0,l2], [0,0,-1,0], [1,0,0,l1+l3], [0,0,0,1]]))
```

---
Cree una función que dados los angulos del manipulador devuelva la transformación total del manipulador (ayudese de la función creada en el segundo problema).

<!-- grade_id: cell-67941966e2bb0f7f, points: 0, grade: false, solution: true -->
```python
def transformacion_PUMA(q1, q2, q3):
    from sympy import pi, var
    var("l1:4")
    ### BEGIN SOLUTION
    A1 = DH_simbolico(0,l1,pi/2, q1)
    A2 = DH_simbolico(l2, 0, 0, q2)
    A3 = DH_simbolico(l3, 0, 0, q3)
    return A1*A2*A3
    ### END SOLUTION
```

<!-- grade_id: cell-1360716371127399, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
from sympy import pi, var, Matrix
var("l1:4")
assert_equal(transformacion_PUMA(0, 0, 0), Matrix([[1,0,0,l2+l3], [0,0,-1,0], [0,1,0,l1], [0,0,0,1]]))
assert_equal(transformacion_PUMA(pi/2, 0, 0), Matrix([[0,0,1,0], [1,0,0,l2+l3], [0,1,0,l1], [0,0,0,1]]))
assert_equal(transformacion_PUMA(0, pi/2, 0), Matrix([[0,-1,0,0], [0,0,-1,0], [1,0,0,l1+l2+l3], [0,0,0,1]]))
assert_equal(transformacion_PUMA(0, 0, pi/2), Matrix([[0,-1,0,l2], [0,0,-1,0], [1,0,0,l1+l3], [0,0,0,1]]))
```

---
Cree una función que dados los angulos del manipulador, grafique las posiciones de los eslabones del manipulador del primer punto (ayudese de las funciones creadas en el primer y segundo problemas, modificadas ligeramente para aceptar matrices numéricas, así como la función creada en la práctica anterior para la graficación de un sistema robótico).

<!-- grade_id: cell-d9e16df1267dfeb6, points: 0, grade: false, solution: true -->
```python
def DH_numerico(a, d, α, θ):
    ### BEGIN SOLUTION
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), -sin(θ)*cos(α), sin(θ)*sin(α), a*cos(θ)],
                 [sin(θ), cos(θ)*cos(α), -cos(θ)*sin(α), a*sin(θ)],
                 [0, sin(α), cos(α), d],
                 [0, 0, 0, 1]])
    return A
    ### END SOLUTION
    
def cinematica_PUMA(q1, q2, q3):
    # Considere que las longitudes son todas iguales a 1
    l1, l2, l3 = 1, 1, 1
    from numpy import pi
    ### BEGIN SOLUTION
    A1 = DH_numerico(0,l1,pi/2, q1)
    A2 = DH_numerico(l2, 0, 0, q2)
    A3 = DH_numerico(l3, 0, 0, q3)
    ### END SOLUTION
    return A1, A2, A3

def grafica_PUMA(q1, q2, q3):
    from numpy import matrix
    ### BEGIN SOLUTION HERE
    A1, A2, A3 = cinematica_PUMA(q1, q2, q3)
    
    H1 = A1
    H2 = A1*A2
    H3 = A1*A2*A3
    
    o0 = matrix([[0],[0],[0],[1]])
    o1, o2, o3 = H1*o0, H2*o0, H3*o0
    
    xs = [o0.item(0), o1.item(0), o2.item(0), o3.item(0)]
    ys = [o0.item(1), o1.item(1), o2.item(1), o3.item(1)]
    zs = [o0.item(2), o1.item(2), o2.item(2), o3.item(2)]
    ### END SOLUTION HERE
    fig = figure(figsize=(8, 8))
    ax = fig.add_subplot(111, projection='3d')

    ax.plot(xs, ys, zs, "-o")
    
    ax.set_xlim(-1.1, 1.1)
    ax.set_ylim(-1.1, 1.1)
    ax.set_zlim(-0.1, 2.1)
    return ax
```

<!-- grade_id: cell-4306e8821b779c0e, points: 3, grade: true, solution: false -->
```python
%matplotlib inline
from matplotlib.pyplot import figure, plot, style
from mpl_toolkits.mplot3d import Axes3D
style.use("ggplot")
from numpy.testing import assert_allclose
from numpy import array
ax = grafica_PUMA(0,0.5,0.5)
ls = ax.get_lines()
assert_allclose(ls[0].get_xdata(), array([0, 0, 0.8775, 1.417885]), rtol=1e-01, atol=1e-01)
assert_allclose(ls[0].get_ydata(), array([-0.0384900179, 0, 0.00915,  0.03809]), rtol=1e-01, atol=1e-01)
```

---
Utilice la función ```interact``` para manipular la posición del manipulador, de tal manera que su posición sea aproximadamente $q_1=0.6rad$, $q_2=0.2rad$ y $q_3 = -0.8rad$

<!-- grade_id: cell-b66fd1a7b96109ff, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
from matplotlib.pyplot import figure, plot, style
from mpl_toolkits.mplot3d import Axes3D
style.use("ggplot")
from ipywidgets import interact
from numpy import pi
τ = 2*pi
### BEGIN SOLUTION
i1 = interact(grafica_PUMA, q1=(-τ, τ), q2=(-τ, τ), q3=(-1.0, 1.0))
### END SOLUTION
```

<!-- grade_id: cell-f447987899e058b0, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal
from numpy import pi
τ = 2*pi
### BEGIN HIDDEN TESTS
q1_slider = i1.widget.children[0]
q2_slider = i1.widget.children[1]
q3_slider = i1.widget.children[2]
assert_almost_equal(q1_slider.value, 0.6, 1)
assert_almost_equal(q2_slider.value, 0.2, 1)
assert_almost_equal(q3_slider.value, -0.8, 1)
### END HIDDEN TESTS
```