# Problemas

---
Implemente el código para las funciones que devuelvan matrices de rotación en $x$ y $y$.

<!-- grade_id: cell-7ee04b3095e7ec0a, points: 0, grade: false, solution: true -->
```python
def rotacion_x(θ):
    # escribe tu codigo aqui
    
    
def rotacion_y(θ):
    # escribe tu codigo aqui
    
```

<!-- grade_id: cell-b149256339fe92c4, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix
assert_allclose(rotacion_x(0), eye(4))
assert_allclose(rotacion_x(pi), matrix([[1, 0, 0, 0], [0, -1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(rotacion_y(0), eye(4))
assert_allclose(rotacion_y(pi), matrix([[-1, 0, 0, 0], [0, 1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
```

---
Implemente el código para las funciones que devuelvan matrices de traslación en $x$, $y$ y $z$.

<!-- grade_id: cell-90b53718156230fc, points: 0, grade: false, solution: true -->
```python
def traslacion_x(x):
    # escribe tu codigo aqui
    

def traslacion_y(y):
    # escribe tu codigo aqui
    
    
def traslacion_z(z):
    # escribe tu codigo aqui
    
```

<!-- grade_id: cell-2e12643fb3f77fcf, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
assert_allclose(traslacion_x(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_x(1), matrix([[1, 0, 0, 1], [0, 1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(1), matrix([[1, 0, 0, 0], [0, 1, 0, 1], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(1), matrix([[1, 0, 0, 0], [0, 1, 0, 0], [0, 0, 1, 1], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
```

---
Implemente una cadena cinemática que describa la transformación hecha por una rotacion en $z$ de $30^o$, una rotación en $y$ de $50^o$ y una traslación en $z$ de $1m$, y guardela en la variable ```T```.

<!-- grade_id: cell-3ca3f60003bde0c7, points: 0, grade: false, solution: true -->
```python
from numpy import pi
τ = 2*pi
# escribe tu codigo aqui

```

<!-- grade_id: cell-ff2321bf85e91a73, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
```

---
Sabemos que un pendulo doble puede ser descrito por una cadena cinemática; implemente una función que tome como argumento los parametros de rotación y traslación de cada uno de sus ejes, y devuelva la posición del actuador final.

<!-- grade_id: cell-ab313b08a490f26b, points: 0, grade: false, solution: true -->
```python
def pendulo_doble(q1, q2, l1, l2):
    # escribe tu codigo aqui
    
```

<!-- grade_id: cell-4c34e757c54ce22a, points: 3, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
assert_allclose(pendulo_doble(0,0,1,1), matrix([[1, 0, 0, 2], [0, 1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(pendulo_doble(pi/2,pi/2,1,1), matrix([[-1, 0, 0, -1], [0, -1, 0, 1], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
```