# Problemas

---
Implemente el código para las funciones que devuelvan matrices de rotación en $x$ y $y$.

<!-- grade_id: cell-7ee04b3095e7ec0a, points: 0, grade: false, solution: true -->
```python
def rotacion_x(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, cos(θ), -sin(θ), 0],
                [0, sin(θ), cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
    
def rotacion_y(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), 0, sin(θ), 0],
                [0, 1, 0, 0],
                [-sin(θ), 0, cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
```

<!-- grade_id: cell-b149256339fe92c4, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix
assert_allclose(rotacion_x(0), eye(4))
assert_allclose(rotacion_x(pi), matrix([[1, 0, 0, 0], [0, -1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(rotacion_y(0), eye(4))
assert_allclose(rotacion_y(pi), matrix([[-1, 0, 0, 0], [0, 1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
### BEGIN HIDDEN TESTS
def rotacion_prueba_x(θ):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, cos(θ), -sin(θ), 0],
                [0, sin(θ), cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    
def rotacion_prueba_y(θ):
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), 0, sin(θ), 0],
                [0, 1, 0, 0],
                [-sin(θ), 0, cos(θ), 0],
                [0, 0, 0, 1]])
    return A

assert_allclose(rotacion_x(pi/6), rotacion_prueba_x(pi/6), rtol=1e-05, atol=1e-5)
assert_allclose(rotacion_y(pi/6), rotacion_prueba_y(pi/6), rtol=1e-05, atol=1e-5)
### END HIDDEN TESTS
```

---
Implemente el código para las funciones que devuelvan matrices de traslación en $x$, $y$ y $z$.

<!-- grade_id: cell-90b53718156230fc, points: 0, grade: false, solution: true -->
```python
def traslacion_x(x):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, x],
                [0, 1, 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE

def traslacion_y(y):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, y],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
    
def traslacion_z(z):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, 0],
                [0, 0, 1, z],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
```

<!-- grade_id: cell-2e12643fb3f77fcf, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
assert_allclose(traslacion_x(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_x(1), matrix([[1, 0, 0, 1], [0, 1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(1), matrix([[1, 0, 0, 0], [0, 1, 0, 1], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(1), matrix([[1, 0, 0, 0], [0, 1, 0, 0], [0, 0, 1, 1], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
### BEGIN HIDDEN TESTS
def traslacion_prueba_x(x):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, x],
                [0, 1, 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A

def traslacion_prueba_y(y):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, y],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    
def traslacion_prueba_z(z):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, 0],
                [0, 0, 1, z],
                [0, 0, 0, 1]])
    return A

assert_allclose(traslacion_x(-1), traslacion_prueba_x(-1), rtol=1e-05, atol=1e-5)
assert_allclose(traslacion_y(-1), traslacion_prueba_y(-1), rtol=1e-05, atol=1e-5)
assert_allclose(traslacion_z(-1), traslacion_prueba_z(-1), rtol=1e-05, atol=1e-5)
### END HIDDEN TESTS
```

---
Implemente una cadena cinemática que describa la transformación hecha por una rotacion en $z$ de $30^o$, una rotación en $y$ de $50^o$ y una traslación en $z$ de $1m$, y guardela en la variable ```T```.

<!-- grade_id: cell-3ca3f60003bde0c7, points: 0, grade: false, solution: true -->
```python
from numpy import pi
τ = 2*pi
### BEGIN SOLUTION HERE
def rotacion_z(θ):
    '''
    Esta función devuelve una matriz de transformación con valores numéricos,
    correspondientes a una rotación alrededor del eje z.
    '''
    from numpy import sin, cos
    A = matrix([[cos(θ), -sin(θ), 0, 0],
                [sin(θ), cos(θ), 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A

T = traslacion_z(1)*rotacion_y(τ*50/360)*rotacion_z(τ*30/360)
### END SOLUTION HERE
```

<!-- grade_id: cell-ff2321bf85e91a73, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
### BEGIN HIDDEN TESTS
T_prueba = traslacion_z(1)*rotacion_y(τ*50/360)*rotacion_z(τ*30/360)
assert_allclose(T, T_prueba, rtol=1e-05, atol=1e-05)
### END HIDDEN TESTS
```

---
Sabemos que un pendulo doble puede ser descrito por una cadena cinemática; implemente una función que tome como argumento los parametros de rotación y traslación de cada uno de sus ejes, y devuelva la posición del actuador final.

<!-- grade_id: cell-ab313b08a490f26b, points: 0, grade: false, solution: true -->
```python
def pendulo_doble(q1, q2, l1, l2):
    ### BEGIN SOLUTION HERE
    R1 = rotacion_z(q1)
    T1 = traslacion_x(l1)
    R2 = rotacion_z(q2)
    T2 = traslacion_x(l2)
    return R1*T1*R2*T2
    ### END SOLUTION HERE
```

<!-- grade_id: cell-4c34e757c54ce22a, points: 3, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
assert_allclose(pendulo_doble(0,0,1,1), matrix([[1, 0, 0, 2], [0, 1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(pendulo_doble(pi/2,pi/2,1,1), matrix([[-1, 0, 0, -1], [0, -1, 0, 1], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
### BEGIN HIDDEN TESTS
def pendulo_doble_prueba(q1, q2, l1, l2):
    R1 = rotacion_z(q1)
    T1 = traslacion_x(l1)
    R2 = rotacion_z(q2)
    T2 = traslacion_x(l2)
    return R1*T1*R2*T2
assert_allclose(pendulo_doble(pi/6, pi/6, 1, 1), pendulo_doble_prueba(pi/6, pi/6, 1, 1), rtol=1e-05, atol=1e-05)
### END HIDDEN TESTS
```