# Problemas

---
Implemente el código para las funciones que devuelvan matrices de rotación y traslación necesarias para implementar la cinemática directa de un manipulador Stanford.

<!-- grade_id: cell-a454ef709499bb24, points: 0, grade: false, solution: true -->
```python
def rotacion_x(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, cos(θ), -sin(θ), 0],
                [0, sin(θ), cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
    
def rotacion_y(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), 0, sin(θ), 0],
                [0, 1, 0, 0],
                [-sin(θ), 0, cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE

def rotacion_z(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), -sin(θ), 0, 0],
                [sin(θ), cos(θ), 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE

def traslacion_x(x):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, x],
                [0, 1, 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE

def traslacion_y(y):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, y],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
    
def traslacion_z(z):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, 0],
                [0, 0, 1, z],
                [0, 0, 0, 1]])
    return A
    ### END SOLUTION HERE
```

<!-- grade_id: cell-f7e193e853d79e64, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
assert_allclose(rotacion_x(0), eye(4))
assert_allclose(rotacion_x(pi), matrix([[1, 0, 0, 0], [0, -1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(rotacion_y(0), eye(4))
assert_allclose(rotacion_y(pi), matrix([[-1, 0, 0, 0], [0, 1, 0, 0], [0, 0, -1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(rotacion_z(0), eye(4))
assert_allclose(rotacion_z(pi), matrix([[-1, 0, 0, 0], [0, -1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_x(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_x(1), matrix([[1, 0, 0, 1], [0, 1, 0, 0], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_y(1), matrix([[1, 0, 0, 0], [0, 1, 0, 1], [0, 0, 1, 0], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(0), eye(4), rtol=1e-05, atol=1e-05)
assert_allclose(traslacion_z(1), matrix([[1, 0, 0, 0], [0, 1, 0, 0], [0, 0, 1, 1], [0, 0, 0, 1]]), rtol=1e-05, atol=1e-05)
### BEGIN HIDDEN TESTS
def rotacion_prueba_x(θ):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, cos(θ), -sin(θ), 0],
                [0, sin(θ), cos(θ), 0],
                [0, 0, 0, 1]])
    return A
    
def rotacion_prueba_y(θ):
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), 0, sin(θ), 0],
                [0, 1, 0, 0],
                [-sin(θ), 0, cos(θ), 0],
                [0, 0, 0, 1]])
    return A

def rotacion_prueba_z(θ):
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), -sin(θ), 0, 0],
                [sin(θ), cos(θ), 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A

def traslacion_prueba_x(x):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, x],
                [0, 1, 0, 0],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A

def traslacion_prueba_y(y):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, y],
                [0, 0, 1, 0],
                [0, 0, 0, 1]])
    return A
    
def traslacion_prueba_z(z):
    from numpy import matrix, sin, cos
    A = matrix([[1, 0, 0, 0],
                [0, 1, 0, 0],
                [0, 0, 1, z],
                [0, 0, 0, 1]])
    return A

assert_allclose(rotacion_x(pi/6), rotacion_prueba_x(pi/6), rtol=1e-05, atol=1e-5)
assert_allclose(rotacion_y(pi/6), rotacion_prueba_y(pi/6), rtol=1e-05, atol=1e-5)
assert_allclose(rotacion_z(pi/6), rotacion_prueba_z(pi/6), rtol=1e-05, atol=1e-5)
assert_allclose(traslacion_x(-1), traslacion_prueba_x(-1), rtol=1e-05, atol=1e-5)
assert_allclose(traslacion_y(-1), traslacion_prueba_y(-1), rtol=1e-05, atol=1e-5)
assert_allclose(traslacion_z(-1), traslacion_prueba_z(-1), rtol=1e-05, atol=1e-5)
### END HIDDEN TESTS
```

---
Implemente una función que devuelva las matrices de transformación homogeneas, que describen la cinemática directa de un manipulador Stanford.

<!-- grade_id: cell-828354d4c89a419a, points: 0, grade: false, solution: true -->
```python
def cinematica_stanford(q1, q2, q3):
    # Tomamos en cuenta que la longitudes son 1
    l1, l2 = 1, 1
    ### BEGIN SOLUTION HERE
    A1 = rotacion_z(q1)*traslacion_z(l1)
    A2 = rotacion_x(q2)
    A3 = traslacion_y(l2 + q3)
    ### END SOLUTION HERE
    return A1, A2, A3
```

```python
cinematica_stanford(0,0,0)
```

<!-- grade_id: cell-54080d520d7bdca1, points: 3, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import eye, pi, matrix, array
A1, A2, A3 = cinematica_stanford(0, 0, 0)
assert_allclose(A1, matrix([[1,0,0,0], [0,1,0,0], [0,0,1,1], [0,0,0,1]]), rtol=1e-05, atol=1e-5)
assert_allclose(A2, eye(4), rtol=1e-05, atol=1e-5)
assert_allclose(A3, matrix([[1,0,0,0], [0,1,0,1], [0,0,1,0], [0,0,0,1]]), rtol=1e-05, atol=1e-5)
### BEGIN HIDDEN TESTS
def cinematica_stanford_prueba(q1, q2, q3):
    # Tomamos en cuenta que la longitudes son 1
    l1, l2 = 1, 1
    A1 = rotacion_z(q1)*traslacion_z(l1)
    A2 = rotacion_x(q2)
    A3 = traslacion_y(l2 + q3)
    return A1, A2, A3

assert_allclose(cinematica_stanford(0,0,0), cinematica_stanford_prueba(0,0,0), rtol=1e-05, atol=1e-5)
assert_allclose(cinematica_stanford(0,0,1), cinematica_stanford_prueba(0,0,1), rtol=1e-05, atol=1e-5)
assert_allclose(cinematica_stanford(0,1,0), cinematica_stanford_prueba(0,1,0), rtol=1e-05, atol=1e-5)
assert_allclose(cinematica_stanford(1,0,0), cinematica_stanford_prueba(1,0,0), rtol=1e-05, atol=1e-5)
assert_allclose(cinematica_stanford(1,1,1), cinematica_stanford_prueba(1,1,1), rtol=1e-05, atol=1e-5)
### END HIDDEN TESTS
```

---
Cree una función que tome como argumentos los valores de posición (angular o lineal) de los grados de libertad de un manipulador Stanford, calcule la cinemática directa y grafique en tres dimensiones un esquema de los eslabones del manipulador.

<!-- grade_id: cell-ac99637b933b8361, points: 0, grade: false, solution: true -->
```python
def grafica_stanford(q1, q2, q3):
    from numpy import matrix
    ### BEGIN SOLUTION HERE
    A1, A2, A3 = cinematica_stanford(q1, q2, q3)
    
    H1 = A1
    H2 = A1*A2
    H3 = A1*A2*A3
    
    o0 = matrix([[0],[0],[0],[1]])
    o1, o2, o3 = H1*o0, H2*o0, H3*o0
    
    xs = [o0.item(0), o1.item(0), o2.item(0), o3.item(0)]
    ys = [o0.item(1), o1.item(1), o2.item(1), o3.item(1)]
    zs = [o0.item(2), o1.item(2), o2.item(2), o3.item(2)]
    ### END SOLUTION HERE
    fig = figure(figsize=(8, 8))
    ax = fig.add_subplot(111, projection='3d')

    ax.plot(xs, ys, zs, "-o")
    
    ax.set_xlim(-1.1, 1.1)
    ax.set_ylim(-1.1, 1.1)
    ax.set_zlim(-0.1, 2.1)
    return ax
```

<!-- grade_id: cell-6cef9847505a66b2, points: 2, grade: true, solution: false -->
```python
%matplotlib inline
from matplotlib.pyplot import figure, plot, style
from mpl_toolkits.mplot3d import Axes3D
style.use("ggplot")
from numpy.testing import assert_allclose
ax = grafica_stanford(0,0.5,0.5)
ls = ax.get_lines()
assert_allclose(ls[0].get_xdata(), array([0, 0, 0,  0.0290876490]), rtol=1e-01, atol=1e-01)
assert_allclose(ls[0].get_ydata(), array([-0.0384900179, 0, 0,  1.316]), rtol=1e-01, atol=1e-01)
```

---
Utilice la función ```interact``` para controlar el movimiento de un manipulador Stanford y posicione al manipulador aproximadamente en $q_1 = 30^o$, $q_2 = 20^o$ y $q_3 = 0.2m$

<!-- grade_id: cell-f6b9ae944f1f3d8e, points: 0, grade: false, solution: true -->
```python
from ipywidgets import interact
from numpy import pi
τ = 2*pi
### BEGIN SOLUTION
i1 = interact(grafica_stanford, q1=(0, τ), q2=(0, τ), q3=(0, 1.0))
### END SOLUTION
```

<!-- grade_id: cell-ef44fd375da5807a, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal
from numpy import pi
τ = 2*pi
### BEGIN HIDDEN TESTS
q1_slider = i1.widget.children[0]
q2_slider = i1.widget.children[1]
q3_slider = i1.widget.children[2]
assert_almost_equal(q1_slider.value, τ/12, 1)
assert_almost_equal(q2_slider.value, τ/18, 1)
assert_almost_equal(q3_slider.value, 0.2, 1)
### END HIDDEN TESTS
```